# Python Advanced

This is autograded assignment. After each problem there is a test cell (with `assert`'s). If it runs correctly, test is passed, if it gives error, something should be fixed. Implement the functions, click "Validate" button, then submit for grading. All problems weight 1 point except where stated otherwise.

## Matrix

Implement class <b>Matrix</b> and exception <b>MatrixSizeError</b>.

Class <b>Matrix</b> has methods:

1. `__init__`, `__str__`. Method `__init__` accepts a list of lists. Internal lists are the same size. The method `__str__` should return a string in a special form. Examples: matrix [[1, 2, 3], [7, 8, 9]] → '1\t2\t3\n7\t8\t9'. Separate lines by new line, elements in one line by tabulation;
1. `__eq__`, `size`. Method `size` should return tuple with 2 elements - (number of rows, number of columns); Method `__eq__` should return True if two matrices are equal (all elements equal), False otherwise;
1. `__add__`, `__sub__` − implement operators + and − for matrices. If the sizes of matrices are not suitable for these operations, throw an exception <b>MatrixSizeError</b>;
1. `__mul__` − multiplying a matrix by a matrix, when multiplying by another data type, throw an exception <b>TypeError</b>. If the sizes of matrices are not suitable for multiplication, throw an exception <b>MatrixSizeError</b>;
1. `transpose` − return a new matrix that is transposed to the current matrix;
1. for square matrix implement methods `tr`(trace) and `det`(determinant, recursively is allowed). For non-square matrix throw an exception <b>MatrixSizeError</b>

- If the argument type is not allowed raise TypeError
- Methods `__add__`, `__sub__` and `__mul__` have to return new Matrix objects. Don't modify self object, methods `__iadd__`, `__isub__` and `__imul__` are used for this purpose (they don't need to be implemented in this task)
- What is '-> "Matrix"' or '-> bool'? https://docs.python.org/3/library/typing.html


### Write the solution

```python
class MatrixSizeError(Exception):
    """
    Exception class for matrix size errors
    """
    pass

class Matrix:
    # Part 1
    def __init__(self, matrix):
        """
        Matrix consturctor
        :param matrix: list of lists
        """
        self.matrix = [list(e) for e in matrix]
    
    def __str__(self) -> str:
        """
        This method returns a string representation of the matrix
        :returns: string
        """
        return "".join(map(lambda e: '\t'.join(map(str, e)) + '\n', self.matrix)).rstrip('\n')

    def __eq__(self, other) -> bool:
        """
        This method checks whether matrices are equal
        :param other: Matrix
        :raises: TypeError, if non-matrix object is passed as the argument
        :returns: bool
        """
        if isinstance(other, Matrix):
            return self.matrix == other.matrix
        else:
            raise TypeError("Cannot compare Matrix with non-Matrix obj")
        
    def size(self) -> tuple:
        """
        This method calculates the matrix size
        :returns: tuple, matrix size
        """
        return (len(self.matrix), len(self.matrix[0]))

    def __add__(self, other: "Matrix") -> "Matrix":
        """
        This method adds two matrices
        :param other: Matrix
        :raises: TypeError, if non-matrix object is passed as the argument
        :raises: ValueError, if matrices have different sizes
        :returns: Matrix sum
        """
        if isinstance(other, Matrix):
            if self.size() == other.size():
                return Matrix([[self.matrix[i][j] + other.matrix[i][j] for j in range(len(self.matrix[0]))] for i in range(len(self.matrix))])
            else:
                raise MatrixSizeError("Matrices have different sizes")
        else:
            raise TypeError("Cannot add Matrix with non-Matrix obj")

    def __sub__(self, other: "Matrix") -> "Matrix":
        """
        This method substracts two matrices
        :param other: Matrix
        :raises: TypeError, if non-matrix object is passed as the argument
        :raises: ValueError, if matrices have different sizes
        :returns: Matrix difference
        """
        if isinstance(other, Matrix):
            if self.size() == other.size():
                return Matrix([[self.matrix[i][j] - other.matrix[i][j] for j in range(len(self.matrix[0]))] for i in range(len(self.matrix))])
            else:
                raise MatrixSizeError("Matrices have different sizes")
        else:
            raise TypeError("Cannot subtract Matrix with non-Matrix obj")

    def __mul__(self, other: "Matrix") -> "Matrix":
        """
        This method multiplies two matrices
        :param other: Matrix
        :raises: TypeError, if non-matrix object is passed as the argument
        :raises: ValueError, if matrices have different sizes
        :returns: Matrix product
        """
        if isinstance(other, Matrix):
            if self.size()[1] == other.size()[0]:
                return Matrix([[sum([self.matrix[i][k] * other.matrix[k][j] for k in range(len(self.matrix[0]))]) for j in range(len(other.matrix[0]))] for i in range(len(self.matrix))])
            else:
                raise MatrixSizeError("Matrices have different sizes")
        else:
            raise TypeError("Cannot multiply Matrix with non-Matrix obj")

    def transpose(self) -> "Matrix":
        """
        This method transposes the matrix
        :returns: transposed matrix
        """
        transposed = [[self.matrix[j][i] for j in range(len(self.matrix))] for i in range(len(self.matrix[0]))]
        return Matrix(transposed)

    def tr(self) -> float:
        """
        This method calculates the trace of the matrix
        :raises: MatrixSizeError, if matrix is non-square
        :returns: trace
        """
        if self.size()[0] == self.size()[1]:
            return sum([self.matrix[i][i] for i in range(len(self.matrix))])
        else:
            raise MatrixSizeError("Matrix is non-square")
    
    def det(self) -> float:
        """
        This method calculates the determinant of the matrix
        :raises: MatrixSizeError, if matrix is non-square
        :returns: determinant
        """
        if len(self.matrix) == 1:
            return self.matrix[0][0]
        elif self.size()[0] == self.size()[1]:
            return sum([(-1) ** i * self.matrix[0][i] * Matrix([[self.matrix[j][k] for k in range(len(self.matrix[0])) if k != i] for j in range(1, len(self.matrix))]).det() for i in range(len(self.matrix[0]))])
        else:
            raise MatrixSizeError("Matrix is non-square")
```

### Run tests

### Part 1

<!-- grade_id: cell-3be9b018b3b49483, points: 1, grade: true, solution: false -->
```python
def check():
    '''
    Part 1
    __init__(self, matrix)
    __str__(self)
    '''

    try:
        matrix_data = [
            [1, 2, 3],
            [4, 5, 6],
        ]
        matrix_str = '1\t2\t3\n4\t5\t6'

        matrix = Matrix(matrix_data)
        if matrix.__str__() != matrix_str:
            return '0\tIncorrect method __str__. Please try again!'

        matrix_data[0][1] = -10
        if matrix.__str__() != matrix_str:
            return '0\tMethod __init__ should copy input data. Use deepcopy.'
    except Exception:
        return '0\tFailed test #1. Please try again!'

    try:
        matrix_data = [[1]]
        matrix_str = '1'

        matrix = Matrix(matrix_data)
        if matrix.__str__() != matrix_str:
            return '0\tIncorrect method __str__. Please try again!'
    except Exception:
        return '0\tFailed test #2. Please try again!'

    try:
        matrix_data = [[1, 3, 5]]
        matrix_str = '1\t3\t5'

        matrix = Matrix(matrix_data)
        if matrix.__str__() != matrix_str:
            return '0\tIncorrect method __str__. Please try again!'
    except Exception:
        return '0\tFailed test #3. Please try again!'

    try:
        matrix_data = [[-10], [90]]
        matrix_str = '-10\n90'

        matrix = Matrix(matrix_data)
        if matrix.__str__() != matrix_str:
            return '0\tIncorrect method __str__. Please try again!'
    except Exception:
        return '0\tFailed test #4. Please try again!'

    try:
        matrix_data = [[78, -10], [90, 15]]
        matrix_str = '78\t-10\n90\t15'

        matrix = Matrix(matrix_data)
        if matrix.__str__() != matrix_str:
            return '0\tIncorrect method __str__. Please try again!'
    except Exception:
        return '0\tFailed test #5. Please try again!'

    try:
        matrix_data = [
            [78, -10, 67, 23, 12, 4],
            [-90, -15, 9, 56, -23, 10000],
            [1, -2, -3, 0, 51, 26]
        ]
        matrix_str = '78\t-10\t67\t23\t12\t4\n-90\t-15\t9\t56\t-23\t10000\n1\t-2\t-3\t0\t51\t26'

        matrix = Matrix(matrix_data)
        if matrix.__str__() != matrix_str:
            return '0\tIncorrect method __str__. Please try again!'
    except Exception:
        return '0\tFailed test #6. Please try again!'

    return '1\tGreat job! You passed all test cases.'


result, message = check().split('\t')
assert result == '1', message
print(message)
```

Output:
```
Great job! You passed all test cases.
```

### Part 2

<!-- grade_id: cell-a3b9c4803bf4f88c, points: 1, grade: true, solution: false -->
```python
def check():
    '''
    Part 2
    __eq__(self, other)
    size(self)
    '''

    # Part 2
    try:
        matrix_data = [
            [1, 2, 3],
            [4, 5, 6],
        ]
        if Matrix(matrix_data).size() != (2, 3): return '0\tIncorrect method size. Please try again!'
        if Matrix(matrix_data) != Matrix(matrix_data): return '0\tIncorrect method __eq__. Please try again!'
        try:
            _ = Matrix(matrix_data) == matrix_data
            return '0\tIncorrect method __eq__. Please try again!'
        except TypeError:
            pass
        except Exception:
            return '0\tIncorrect method __eq__. Please try again!'

        matrix_data = [[1]]
        if Matrix(matrix_data).size() != (1, 1): return '0\tIncorrect method size. Please try again!'
        if Matrix(matrix_data) != Matrix(matrix_data): return '0\tIncorrect method __eq__. Please try again!'
        try:
            _ = Matrix(matrix_data) == 5
            return '0\tIncorrect method __eq__. Please try again!'
        except TypeError:
            pass
        except Exception:
            return '0\tIncorrect method __eq__. Please try again!'

        matrix_data = [[1, 3, 5]]
        if Matrix(matrix_data).size() != (1, 3): return '0\tIncorrect method size. Please try again!'
        if Matrix(matrix_data) != Matrix(matrix_data): return '0\tIncorrect method __eq__. Please try again!'
        try:
            _ = Matrix(matrix_data) == '2345678sdfgh'
            return '0\tIncorrect method __eq__. Please try again!'
        except TypeError:
            pass
        except Exception:
            return '0\tIncorrect method __eq__. Please try again!'

        matrix_data = [[-10], [90]]
        if Matrix(matrix_data).size() != (2, 1): return '0\tIncorrect method size. Please try again!'
        if Matrix(matrix_data) != Matrix(matrix_data): return '0\tIncorrect method __eq__. Please try again!'
        try:
            _ = Matrix(matrix_data) == (1, 3, 4)
            return '0\tIncorrect method __eq__. Please try again!'
        except TypeError:
            pass
        except Exception:
            return '0\tIncorrect method __eq__. Please try again!'

        matrix_data = [[78, -10], [90, 15]]
        if Matrix(matrix_data).size() != (2, 2): return '0\tIncorrect method size. Please try again!'
        if Matrix(matrix_data) != Matrix(matrix_data): return '0\tIncorrect method __eq__. Please try again!'
        try:
            _ = Matrix(matrix_data) == {'hello': 1}
            return '0\tIncorrect method __eq__. Please try again!'
        except TypeError:
            pass
        except Exception:
            return '0\tIncorrect method __eq__. Please try again!'

        matrix_data = [
            [78, -10, 67, 23, 12, 4],
            [-90, -15, 9, 56, -23, 10000],
            [1, -2, -3, 0, 51, 26]
        ]
        if Matrix(matrix_data).size() != (3, 6): return '0\tIncorrect method size. Please try again!'
        if Matrix(matrix_data) != Matrix(matrix_data): return '0\tIncorrect method __eq__. Please try again!'
    except Exception:
        return '0\tFailed test #7. Please try again!'

    return '1\tGreat job! You passed all test cases.'


result, message = check().split('\t')
assert result == '1', message
print(message)
```

Output:
```
Great job! You passed all test cases.
```

### Part 3

<!-- grade_id: cell-23e95a6be2244bf2, points: 1, grade: true, solution: false -->
```python
def check():
    '''
    Part 3
    __add__(self, other)
    __sub__(self, other)
    '''

    # Part 3
    try:
        if not issubclass(MatrixSizeError, Exception):
            return "0\tMatrixSizeError doesn't inherit Exception"

        matrix_data_1 = [[1, 2, 3], [4, 5, 6]]
        matrix_data_2 = [[7, -8, 2], [-14, 17, -8]]
        matrix_data_add = [[8, -6, 5], [-10, 22, -2]]
        matrix_data_sub = [[-6, 10, 1], [18, -12, 14]]
        if Matrix(matrix_data_1) + Matrix(matrix_data_2) != Matrix(matrix_data_add): return '0\tIncorrect method __add__. Please try again!'
        if Matrix(matrix_data_1) - Matrix(matrix_data_2) != Matrix(matrix_data_sub): return '0\tIncorrect method __sub__. Please try again!'

        try:
            _ = Matrix(matrix_data_1) + 5
            return '0\tIncorrect method __add__. Please try again!'
        except TypeError: pass
        except Exception: return '0\tIncorrect method __add__. Please try again!'

        try:
            _ = Matrix(matrix_data_1) - 5
            return '0\tIncorrect method __sub__. Please try again!'
        except TypeError: pass
        except Exception: return '0\tIncorrect method __sub__. Please try again!'

        matrix_data_1 = [[1]]
        matrix_data_2 = [[3]]
        matrix_data_add = [[4]]
        matrix_data_sub = [[-2]]
        if Matrix(matrix_data_1) + Matrix(matrix_data_2) != Matrix(matrix_data_add): return '0\tIncorrect method __add__. Please try again!'
        if Matrix(matrix_data_1) - Matrix(matrix_data_2) != Matrix(matrix_data_sub): return '0\tIncorrect method __sub__. Please try again!'

        try:
            _ = Matrix(matrix_data_1) + list(range(100))
            return '0\tIncorrect method __add__. Please try again!'
        except TypeError: pass
        except Exception: return '0\tIncorrect method __add__. Please try again!'

        try:
            _ = Matrix(matrix_data_1) - list(range(100))
            return '0\tIncorrect method __sub__. Please try again!'
        except TypeError: pass
        except Exception: return '0\tIncorrect method __sub__. Please try again!'

        matrix_data_1 = [[1, 3, 5]]
        matrix_data_2 = [[-7, 8, 0]]
        matrix_data_add = [[-6, 11, 5]]
        matrix_data_sub = [[8, -5, 5]]
        if Matrix(matrix_data_1) + Matrix(matrix_data_2) != Matrix(matrix_data_add): return '0\tIncorrect method __add__. Please try again!'
        if Matrix(matrix_data_1) - Matrix(matrix_data_2) != Matrix(matrix_data_sub): return '0\tIncorrect method __sub__. Please try again!'

        try:
            _ = Matrix(matrix_data_1) + {'a': 1, 'b': 2}
            return '0\tIncorrect method __add__. Please try again!'
        except TypeError: pass
        except Exception: return '0\tIncorrect method __add__. Please try again!'

        try:
            _ = Matrix(matrix_data_1) - {'a': 1, 'b': 2}
            return '0\tIncorrect method __sub__. Please try again!'
        except TypeError: pass
        except Exception: return '0\tIncorrect method __sub__. Please try again!'

        matrix_data_1 = [[-10], [90]]
        matrix_data_2 = [[10], [-90]]
        matrix_data_add = [[0], [0]]
        matrix_data_sub = [[-20], [180]]
        if Matrix(matrix_data_1) + Matrix(matrix_data_2) != Matrix(matrix_data_add): return '0\tIncorrect method __add__. Please try again!'
        if Matrix(matrix_data_1) - Matrix(matrix_data_2) != Matrix(matrix_data_sub): return '0\tIncorrect method __sub__. Please try again!'

        matrix_1 = Matrix([[-10], [90]])
        matrix_2 = Matrix([[1], [3], [5]])
        try:
            _ = matrix_1 + matrix_2
            return '0\tIncorrect method __add__. Please try again!'
        except MatrixSizeError: pass
        except Exception: return '0\tIncorrect method __add__. Please try again!'

        try:
            _ = matrix_1 - matrix_2
            return '0\tIncorrect method __sub__. Please try again!'
        except MatrixSizeError: pass
        except Exception: return '0\tIncorrect method __sub__. Please try again!'

        matrix_1 = Matrix([[-10], [90]])
        matrix_2 = Matrix([[1, 3]])
        try:
            _ = matrix_1 + matrix_2
            return '0\tIncorrect method __add__. Please try again!'
        except MatrixSizeError: pass
        except Exception: return '0\tIncorrect method __add__. Please try again!'

        try:
            _ = matrix_1 - matrix_2
            return '0\tIncorrect method __sub__. Please try again!'
        except MatrixSizeError: pass
        except Exception: return '0\tIncorrect method __sub__. Please try again!'

        matrix_1 = Matrix([[1, 2, 3], [4, 5, 6]])
        matrix_2 = Matrix([[-7, 8, 0]])
        try:
            _ = matrix_1 + matrix_2
            return '0\tIncorrect method __add__. Please try again!'
        except MatrixSizeError: pass
        except Exception: return '0\tIncorrect method __add__. Please try again!'

        try:
            _ = matrix_1 - matrix_2
            return '0\tIncorrect method __sub__. Please try again!'
        except MatrixSizeError: pass
        except Exception: return '0\tIncorrect method __sub__. Please try again!'

    except Exception:
        return '0\tFailed test #8. Please try again!'

    return '1\tGreat job! You passed all test cases.'


result, message = check().split('\t')
assert result == '1', message
print(message)
```

Output:
```
Great job! You passed all test cases.
```

### Part 4

<!-- grade_id: cell-96e6d14da1e1187e, points: 1, grade: true, solution: false -->
```python
def check():
    '''
    Part 4
    __mul__(self, other)
    '''

    # Part 4
    try:
        matrix_1 = Matrix([[1, 2], [4, 5], [7, 8]])
        matrix_2 = Matrix([[-8, 9, 12], [-7, 5, 6]])
        matrix_3 = Matrix([[-22, 19, 24], [-67, 61, 78], [-112, 103, 132]])
        if matrix_1 * matrix_2 != matrix_3:
            return '0\tIncorrect method __mul__. Please try again!'

        matrix_3 = Matrix([[112, 125], [55, 59]])
        if matrix_2 * matrix_1 != matrix_3:
            return '0\tIncorrect method __mul__. Please try again!'

        matrix_1 = Matrix([[6]])
        matrix_2 = Matrix([[7]])
        matrix_3 = Matrix([[42]])
        if matrix_1 * matrix_2 != matrix_3:
            return '0\tIncorrect method __mul__. Please try again!'
        if matrix_2 * matrix_1 != matrix_3:
            return '0\tIncorrect method __mul__. Please try again!'

        matrix_1 = Matrix([[-12345]])
        matrix_2 = Matrix([[0]])
        matrix_3 = Matrix([[0]])
        if matrix_1 * matrix_2 != matrix_3:
            return '0\tIncorrect method __mul__. Please try again!'
        if matrix_2 * matrix_1 != matrix_3:
            return '0\tIncorrect method __mul__. Please try again!'

        matrix_1 = Matrix([[1, 2], [4, 5], [7, 8], [-9, 0]])
        matrix_2 = Matrix([[-8, 9, 12], [-7, 5, 6]])
        matrix_3 = Matrix([[-22, 19, 24], [-67, 61, 78], [-112, 103, 132], [72, -81, -108]])
        if matrix_1 * matrix_2 != matrix_3:
            return '0\tIncorrect method __mul__. Please try again!'

        try:
            _ = matrix_2 * matrix_1
            return '0\tIncorrect method __mul__. Please try again!'
        except MatrixSizeError: pass
        except Exception: return '0\tIncorrect method __mul__. Please try again!'

        try:
            _ = matrix_1 * 5
            return '0\tIncorrect method __mul__. Please try again!'
        except TypeError: pass
        except Exception: return '0\tIncorrect method __mul__. Please try again!'

        try:
            _ = matrix_2 * [5, 4, 6, 1]
            return '0\tIncorrect method __mul__. Please try again!'
        except TypeError: pass
        except Exception: return '0\tIncorrect method __mul__. Please try again!'

        try:
            _ = matrix_3 * {'a': 1, 'b': 2}
            return '0\tIncorrect method __mul__. Please try again!'
        except TypeError: pass
        except Exception: return '0\tIncorrect method __mul__. Please try again!'

        try:
            _ = Matrix([[-12345]]) * Matrix([[-12345], [-12345]])
            return '0\tIncorrect method __mul__. Please try again!'
        except MatrixSizeError: pass
        except Exception: return '0\tIncorrect method __mul__. Please try again!'

    except Exception:
        return '0\tFailed test #9. Please try again!'

    return '1\tGreat job! You passed all test cases.'


result, message = check().split('\t')
assert result == '1', message
print(message)
```

Output:
```
Great job! You passed all test cases.
```

### Part 5

<!-- grade_id: cell-109bad3d3779b79c, points: 1, grade: true, solution: false -->
```python
def check():
    '''
    Part 5
    transpose(self)
    '''

    # Part 5
    try:
        matrix_1 = Matrix([[1, 2], [4, 5], [7, 8]])
        matrix_2 = Matrix([[1, 4, 7], [2, 5, 8]])
        if matrix_1.transpose() != matrix_2:
            return '0\tIncorrect method transpose. Please try again!'

        if matrix_2.transpose() != matrix_1:
            return '0\tIncorrect method transpose. Please try again!'

        if matrix_1.transpose().transpose() != matrix_1:
            return '0\tIncorrect method transpose. Please try again!'

        if matrix_2.transpose().transpose() != matrix_2:
            return '0\tIncorrect method transpose. Please try again!'

        matrix_1 = Matrix([[1, 2], [2, 1]])
        if matrix_1.transpose() != matrix_1:
            return '0\tIncorrect method transpose. Please try again!'

        if matrix_1.transpose().transpose() != matrix_1:
            return '0\tIncorrect method transpose. Please try again!'

        matrix_1 = Matrix([[100]])
        if matrix_1.transpose() != matrix_1:
            return '0\tIncorrect method transpose. Please try again!'

        if matrix_1.transpose().transpose() != matrix_1:
            return '0\tIncorrect method transpose. Please try again!'

    except Exception:
        return '0\tFailed test #10. Please try again!'

    return '1\tGreat job! You passed all test cases.'


result, message = check().split('\t')
assert result == '1', message
print(message)
```

Output:
```
Great job! You passed all test cases.
```

### Part 6

<!-- grade_id: cell-bb0492a96f34ddd5, points: 1, grade: true, solution: false -->
```python
def check():
    '''
    Part 6
    tr(self)
    det(self)
    '''

    # Part 6
    try:
        # tr
        matrix = Matrix([[1, 2], [4, 5], [7, 8]])
        try:
            _ = matrix.tr()
            return '0\tIncorrect method tr. Please try again!'
        except MatrixSizeError: pass
        except Exception: return '0\tIncorrect method tr. Please try again!'

        try:
            _ = matrix.transpose().tr()
            return '0\tIncorrect method tr. Please try again!'
        except MatrixSizeError: pass
        except Exception: return '0\tIncorrect method tr. Please try again!'

        matrix = Matrix([[1, 2], [4, 5]])
        if matrix.tr() != 6:
            return '0\tIncorrect method transpose. Please try again!'

        if matrix.transpose().tr() != 6:
            return '0\tIncorrect method transpose. Please try again!'

        matrix = Matrix([[15]])
        if matrix.tr() != 15:
            return '0\tIncorrect method transpose. Please try again!'

        if matrix.transpose().tr() != 15:
            return '0\tIncorrect method transpose. Please try again!'

        matrix = Matrix([[1, 2, 3], [4, 5, 6], [7, 8, 9]])
        if matrix.tr() != 15:
            return '0\tIncorrect method transpose. Please try again!'

        if matrix.transpose().tr() != 15:
            return '0\tIncorrect method transpose. Please try again!'

        # det
        matrix = Matrix([[1, 2], [4, 5], [7, 8]])
        try:
            _ = matrix.det()
            return '0\tIncorrect method det. Please try again!'
        except MatrixSizeError: pass
        except Exception: return '0\tIncorrect method det. Please try again!'

        matrix = Matrix([[1, 2, 3], [4, 5, 6], [7, 8, 9]])
        if matrix.det() != 0:
            return '0\tIncorrect method det. Please try again!'

        matrix = Matrix([[0]])
        if matrix.det() != 0:
            return '0\tIncorrect method det. Please try again!'

        matrix = Matrix([[7]])
        if matrix.det() != 7:
            return '0\tIncorrect method det. Please try again!'

        matrix = Matrix([[1, 90], [-7, 6]])
        if matrix.det() != 636:
            return '0\tIncorrect method det. Please try again!'

        matrix = Matrix([[98, -4, -7], [0, 23, 87], [-87, 6, 1]])
        if matrix.det() != -32633:
            return '0\tIncorrect method det. Please try again!'

        matrix = Matrix([[3, -9, 0, -6], [7, 4, 1, 8], [9, 12, -21, 8], [9, 7, -5, 2]])
        if matrix.det() != 12816:
            return '0\tIncorrect method det. Please try again!'

    except Exception:
        return '0\tFailed test #11. Please try again!'

    return '1\tGreat job! You passed all test cases.'


result, message = check().split('\t')
assert result == '1', message
print(message)
```

Output:
```
Great job! You passed all test cases.
```